---
title       : Übungsblatt 2
description : Annahmen des einfachen linearen Regressionsmodells, Erwartungswert und Varianz von Zufallsvariablen, Wiederholte Stichproben und Normalverteilung

--- type:NormalExercise lang:r xp:100 skills:1 key:3f15f6c27a
## Aufgabe 1


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
