---
title       : Übungsblatt 2
description : Annahmen des einfachen linearen Regressionsmodells, Erwartungswert und Varianz von Zufallsvariablen, Wiederholte Stichproben und Normalverteilung

--- type:NormalExercise lang:r xp:100 skills:1 key:3f15f6c27a
## Würfelbeispiel - Aufgabe 1
Berechnen Sie mit den Erwartungswert der Zufallsvariable `u1 = Geworfene Augenzahl bei einmaligem Würfeln` und geben Sie dem Erwartungswert den Namen `E.u1`. Definieren Sie zu diesem Zweck im Quelltext-Fenster den Vektor `u1`, welcher die sechs möglichen Ausprägungen der Zufallsvariable `u1` enthält und nutzen Sie den Befehl `mean(u1)`, der
das arithmetische Mittel des Vektors `u1` berechnet. Lassen Sie sich den Wert von `E.u1` anzeigen.

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
